<div class="alert alert-info">
    
➡️ Make sure that you have read the **[rules for hand-in assignments](https://www.ida.liu.se/~TDDE16/exam.en.shtml#handins)** and the **[policy on cheating and plagiarism](https://www.ida.liu.se/~TDDE16/exam.en.shtml#cheating)** before starting with this lab.

➡️ Make sure you fill in any cells (and _only_ those cells) that say **`YOUR CODE HERE`** or **YOUR ANSWER HERE**, and do _not_ modify any of the other cells.

➡️ **Before you submit your lab, make sure everything runs as expected.** For this, _restart the kernel_ and _run all cells_ from top to bottom. In Jupyter Notebook version 7 or higher, you can do this via "Run$\rightarrow$Restart Kernel and Run All Cells..." in the menu (or the "⏩" button in the toolbar).

</div>

# L4: Information extraction

Information extraction (IE) is the task of identifying named entities and semantic relations between these entities in text data. In this lab we will focus on two sub-tasks in IE, **named entity recognition** (identifying mentions of entities) and **entity linking** (matching these mentions to entities in a knowledge base).

```python
# Define some helper functions that are used in this notebook

from IPython.display import display, HTML

def success():
    display(HTML('<div class="alert alert-success"><strong>Solution appears correct!</strong></div>'))
```

## Data set

We start by loading spaCy. However, the data that we will be using has been tokenized following the conventions of the [Penn Treebank](ftp://ftp.cis.upenn.edu/pub/treebank/public_html/tokenization.html), and we need to prevent spaCy from using its own tokenizer on top of this. We therefore override spaCy&rsquo;s tokenizer with the default one that simply splits on whitespace.

```python
import spacy
from spacy.tokenizer import Tokenizer

nlp = spacy.load('en_core_web_sm')
nlp.tokenizer = Tokenizer(nlp.vocab)
```

The main data set for this lab is a collection of news wire articles in which mentions of named entities have been annotated with page names from the [English Wikipedia](https://en.wikipedia.org/wiki/). The next code cell loads the training and the development parts of the data into Pandas data frames.

```python
import bz2
import csv
import pandas as pd
import numpy as np

with bz2.open('ner-train.tsv.bz2', 'rt') as source:
    df_train = pd.read_csv(source, sep='\t', quoting=csv.QUOTE_NONE)

with bz2.open('ner-dev.tsv.bz2', 'rt') as source:
    df_dev = pd.read_csv(source, sep='\t', quoting=csv.QUOTE_NONE)
```

Each row in these two data frames corresponds to one mention of a named entity and has five columns:

1. a unique identifier for the sentence containing the entity mention
2. the pre-tokenized sentence, with tokens separated by spaces
3. the start position of the token span containing the entity mention
4. the end position of the token span (exclusive, as in Python list indexing)
5. the entity label; either a Wikipedia page name or the generic label `--NME--`

The following cell prints the first five samples from the training data:

```python
df_train.head()
```

Output:
```
sentence_id                                          sentence  beg  end   
0    0000-000  EU rejects German call to boycott British lamb .    0    1  \
1    0000-000  EU rejects German call to boycott British lamb .    2    3   
2    0000-000  EU rejects German call to boycott British lamb .    6    7   
3    0000-001                                   Peter Blackburn    0    2   
4    0000-002                               BRUSSELS 1996-08-22    0    1   

            label  
0         --NME--  
1         Germany  
2  United_Kingdom  
3         --NME--  
4        Brussels
```

In this sample, we see that the first sentence is annotated with three entity mentions:

* the span 0–1 &lsquo;EU&rsquo; is annotated as a mention but only labelled with the generic `--NME--`
* the span 2–3 &lsquo;German&rsquo; is annotated with the page [Germany](http://en.wikipedia.org/wiki/Germany)
* the span 6–7 &lsquo;British&rsquo; is annotated with the page [United_Kingdom](http://en.wikipedia.org/wiki/United_Kingdom)

## Problem 1: Evaluation measures

To warm up, we ask you to write code to print the three measures that you will be using for evaluation:

<!-- grade_id: cell-ca0fe4e7e601cd70, points: 0, grade: false, solution: true -->
```python
def evaluation_scores(gold, pred):
    """Print precision, recall, and F1 score.
    
    Arguments:
        gold: The set with the gold-standard values.
        pred: The set with the predicted values.
    
    Returns:
        A tuple or list containing the precision, recall, and F1 values
        (in that order), computed based on the specified sets.
    """
    # YOUR CODE HERE
    precision = len(gold.intersection(pred))/len(pred)
    recall = len(gold.intersection(pred))/len(gold)
    F1= 2/((1/precision)+(1/recall))

    return [precision, recall, F1]
    
    
```

Let's also define a convenience function that prints the scores nicely:

```python
def print_evaluation_scores(scores):
    p, r, f = scores
    print(f"Precision: {p:.3f}, Recall: {r:.3f}, F1: {f:.3f}")
```

### 🤞 Test your code

To test your code, you can run the following cell. This should give you a precision of 60%, a recall of 100%, and an F1-value of 75%.

<!-- grade_id: cell-0de2d4e882169381, points: 1, grade: true, solution: false -->
```python
# Check if the results match what is expected
result = evaluation_scores(set(range(3)), set(range(5)))
assert len(result) == 3, "Should return exactly three scores"
print_evaluation_scores(result)
assert np.isclose(result, (.6, 1.0, .75)).all(), "Should be close to the expected values"
success()
```

Output:
```
Precision: 0.600, Recall: 1.000, F1: 0.750
<IPython.core.display.HTML object>
```

## Problem 2: Span recognition

One of the first tasks that an information extraction system has to solve is to locate and classify (mentions of) named entities, such as persons and organizations. Here we will tackle the simpler task of recognizing **spans** of tokens that contain an entity mention, without the actual entity label.

The English language model in spaCy features a full-fledged [named entity recognizer](https://spacy.io/usage/linguistic-features#named-entities) that identifies a variety of entities, and can be updated with new entity types by the user. Your task in this problem is to evaluate the performance of this component when predicting entity spans in the development data.

### Task 2.1

Start by implementing a generator function that yields the gold-standard spans in a given data frame.  (If you're not familiar with the `yield` keyword in Python, check out [this brief explanation](https://www.nbshare.io/notebook/851988260/Python-Yield/).)

**Hint:** The Pandas method [`itertuples()`](https://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.itertuples.html) is useful when iterating over the rows in a DataFrame.

<!-- grade_id: cell-96b73b0db9000ea3, points: 0, grade: false, solution: true -->
```python
def gold_spans(df):
    """Yield the gold-standard mention spans in a data frame.

    Arguments:
        df: A data frame.

    Yields:
        The gold-standard mention spans in the specified data frame as
        triples consisting of the sentence id, start position, and end
        position of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        yield((row[1],row[3],row[4]))
```

#### 🤞 Test your code

To test your code, you can run the following cell, which counts the spans yielded by your function when called on the development data (there should be 5,917 _unique_ triples), and checks if the first and last yielded triples are included in the results.

<!-- grade_id: cell-8cb736d325aa10ba, points: 1, grade: true, solution: false -->
```python
spans_dev_gold = set(gold_spans(df_dev))
assert len(spans_dev_gold) == 5917, "The number of unique returned triples is not correct."
assert ('0946-000', 2, 3) in spans_dev_gold, "The first expected triple is not included in the results."
assert ('1161-010', 1, 3) in spans_dev_gold, "The last expected triple is not included in the results."
success()
```

Output:
```
<IPython.core.display.HTML object>
```

### Task 2.2

Your next task is to write code that calls spaCy to predict the named entities in the development data.  You should do this in form of a function that works the same as `gold_spans()`, but which returns the spans as predicted by spaCy instead.

```python
# next(gold_spans(df_dev))
```

<!-- grade_id: cell-f5b568125d3e20e9, points: 0, grade: false, solution: true -->
```python
def pred_spans(df):
    """Run and evaluate spaCy's NER.

    Arguments:
        df: A data frame.

    Yields:
        The predicted mention spans in the specified data frame as
        triples consisting of the sentence id, start position, and end
        position of each span.
    """

    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            yield((row[1],ent.start, ent.end)) 
    
```

#### 🤞 Test your code

The following cell runs the prediction and reports the evaluation measures. The expected precision is above 50%, with a recall above 70% and an F1-score around 60%.


<!-- grade_id: cell-a61480e94bbc6fc7, points: 1, grade: true, solution: false -->
```python
spans_dev_pred = set(pred_spans(df_dev))
spans_dev_gold = set(gold_spans(df_dev))
scores = evaluation_scores(spans_dev_gold, spans_dev_pred)
print_evaluation_scores(scores)
assert scores[0] > .50, "Precision should be above 50%."
assert scores[1] > .70, "Recall should be above 70%."
success()
```

Output:
```
Precision: 0.519, Recall: 0.717, F1: 0.602
<IPython.core.display.HTML object>
```

## Problem 3: Error analysis

As you were able to see in Problem&nbsp;2, the span accuracy of the named entity recognizer is far from perfect. In particular, only slightly more than half of the predicted spans are correct according to the gold standard. Your next task is to analyse this result in more detail.

Here is a function that prints the false positives as well as the false negatives spans for a data frame, given a reference set of gold-standard spans and a candidate set of predicted spans.

```python
from collections import defaultdict

def error_report(df, spans_gold, spans_pred):
    false_pos = defaultdict(list)
    for s, b, e in spans_pred - spans_gold:
        false_pos[s].append((b, e))
    false_neg = defaultdict(list)
    for s, b, e in spans_gold - spans_pred:
        false_neg[s].append((b, e))
    for row in df.drop_duplicates('sentence_id').itertuples():
        if row.sentence_id in false_pos or row.sentence_id in false_neg:
            print('Sentence:', row.sentence)
            for b, e in false_pos[row.sentence_id]:
                print('  FP:', ' '.join(row.sentence.split()[b:e]))
            for b, e in false_neg[row.sentence_id]:
                print('  FN:', ' '.join(row.sentence.split()[b:e]))
```

### Task 3.1

1. Use the `error_report()` function above to inspect and analyse the errors that the automated prediction makes. Base your analysis on the first 500 rows of the training data. Can you see any patterns?
2. Summarize your observations in a short text.

<!-- grade_id: cell-a17e92a61aedec84, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
error_report(df_dev.iloc[:500,:],set(gold_spans(df_dev)),set(pred_spans(df_dev)))
```

Output:
```
Sentence: CRICKET - LEICESTERSHIRE TAKE OVER AT TOP AFTER INNINGS VICTORY .
  FN: LEICESTERSHIRE
Sentence: LONDON 1996-08-30
  FP: 1996-08-30
Sentence: West Indian all-rounder Phil Simmons took four for 38 on Friday as Leicestershire beat Somerset by an innings and 39 runs in two days to take over at the head of the county championship .
  FP: Friday
  FP: 38
  FP: two days
  FP: 39
  FP: four
  FN: Somerset
Sentence: Their stay on top , though , may be short-lived as title rivals Essex , Derbyshire and Surrey all closed in on victory while Kent made up for lost time in their rain-affected match against Nottinghamshire .
  FP: Essex , Derbyshire and
  FN: Essex
  FN: Derbyshire
Sentence: After bowling Somerset out for 83 on the opening morning at Grace Road , Leicestershire extended their first innings by 94 runs before being bowled out for 296 with England discard Andy Caddick taking three for 83 .
  FP: 296
  FP: three
  FP: 83
  FP: the opening morning
  FP: 83
  FP: 94
  FP: first

```

<!-- grade_id: cell-e4b523a1fe8201ea, points: 1, grade: true, solution: true -->
**Summary**:

Most of errors belong to FP error. Accoding to the analysis result, we can see that most of false positive errors are about "ORDINAL","DATE","CARDINAL". These elements are not likely to appear in Wikipedia as entities. From this we can conclude that only entities that are relevant to wikipedia are saved in df_dev dataset.

### Task 3.2

Now, use the insights from your error analysis to improve the automated prediction that you implemented in Problem&nbsp;2. While the best way to do this would be to [update spaCy&rsquo;s NER model](https://spacy.io/usage/linguistic-features#updating) using domain-specific training data, for this lab it suffices to write code to post-process the output produced by spaCy. To filter out specific labels it is useful to know the named entity label scheme, which can be found in the [model's documentation](https://spacy.io/models/en#en_core_web_sm). You should be able to improve the F1 score from Problem&nbsp;2 by at least 15 percentage points.

<!-- grade_id: cell-0dc8ed778c02fc71, points: 0, grade: false, solution: true -->
```python
'''
def pred_spans_improved(df):
    """Run and evaluate spaCy's NER, with post-processing to improve the results.

    Arguments:
        df: A data frame.

    Yields:
        The predicted mention spans in the specified data frame as
        triples consisting of the sentence id, start position, and end
        position of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            if ent.label_!= 'ORDINAL'and ent.label_!= 'DATE' and ent.label_ != 'CARDINAL':
                yield((row[1],ent.start, ent.end)) 

                '''
```

Output:
```
'\ndef pred_spans_improved(df):\n    """Run and evaluate spaCy\'s NER, with post-processing to improve the results.\n\n    Arguments:\n        df: A data frame.\n\n    Yields:\n        The predicted mention spans in the specified data frame as\n        triples consisting of the sentence id, start position, and end\n        position of each span.\n    """\n    # YOUR CODE HERE\n    for row in df.itertuples():\n        doc = nlp(row[2])\n        for ent in doc.ents:\n            if ent.label_!= \'ORDINAL\'and ent.label_!= \'DATE\' and ent.label_ != \'CARDINAL\':\n                yield((row[1],ent.start, ent.end)) \n\n                '
```

```python
def pred_spans_improved(df):
    """Run and evaluate spaCy's NER, with post-processing to improve the results.

    Arguments:
        df: A data frame.

    Yields:
        The predicted mention spans in the specified data frame as
        triples consisting of the sentence id, start position, and end
        position of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            if ent.label_!= 'ORDINAL'and ent.label_!= 'DATE' and ent.label_ != 'CARDINAL' and ent.label_ != 'MONEY' and ent.label_ != 'PERCENT' and ent.label_ != 'QUANTITY' and ent.label_ != 'TIME':
                yield((row[1],ent.start, ent.end)) 
```

#### 🤞 Test your code

The following cell reports the evaluation measures from the new function and tests if you achieve the performance goal.

<!-- grade_id: cell-faf957ffa2d4e385, points: 1, grade: true, solution: false -->
```python
scores_old = evaluation_scores(spans_dev_gold, spans_dev_pred)
scores_new = evaluation_scores(spans_dev_gold, set(pred_spans_improved(df_dev)))
print_evaluation_scores(scores_new)
assert scores_new[-1] - scores_old[-1] > .15, "F1-score should improve by at least 15 percentage points."
success()
```

Output:
```
Precision: 0.859, Recall: 0.716, F1: 0.781
<IPython.core.display.HTML object>
```

### Task 3.3

Before moving on, we ask you to store the outputs of the improved named entity recognizer on the development data in a new data frame. This new frame should have the same layout as the original data frame for the development data that you loaded above, but should contain the *predicted* start and end positions for each token span, rather than the gold positions. As the `label` of each span, you can use the special value `--NME--`.

<!-- grade_id: cell-ef6dfcb92ae20097, points: 0, grade: false, solution: true -->
```python
def df_with_pred_spans(df):
    """Make a new DataFrame with *predicted* NER spans.

    Arguments:
        df: A data frame.

    Returns:
        A *new* data frame with the same layout as `df`t contain, buing
        the predicted start and end positions for each token span.
    """
    # YOUR CODE HERE
    new_dataframe = pd.DataFrame(columns=["sentence_id","sentence","beg","end","label"])
    #list =[]
    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            if ent.label_!= 'ORDINAL'and ent.label_!= 'DATE' and ent.label_ != 'CARDINAL' and ent.label_ != 'MONEY' and ent.label_ != 'PERCENT' and ent.label_ != 'QUANTITY' and ent.label_ != 'TIME':
                sentence_id = row[1]
                sentence = row[2]
                beg= ent.start
                end= ent.end
                
                new_row = pd.DataFrame([[sentence_id,sentence,beg,end,"--NNE--"]],
                                       columns=["sentence_id","sentence","beg","end","label"])
                #list.append([sentence_id,sentence,beg,end,"--NNE--"])
                new_dataframe=pd.concat([new_dataframe,new_row])
    return (new_dataframe)




```

#### 🤞 Test your code

Run the following cell to run your function and display the first few lines of the new data frame:

<!-- grade_id: cell-45725442562de561, points: 1, grade: true, solution: false -->
```python
df_dev_pred = df_with_pred_spans(df_dev)
display(df_dev_pred.head())

```

Output:
```
sentence_id                                           sentence beg end   
0    0946-001                                  LONDON 1996-08-30   0   1  \
0    0946-002  West Indian all-rounder Phil Simmons took four...   0   2   
0    0946-002  West Indian all-rounder Phil Simmons took four...   3   5   
0    0946-002  West Indian all-rounder Phil Simmons took four...  12  13   
0    0946-002  West Indian all-rounder Phil Simmons took four...   0   2   

     label  
0  --NNE--  
0  --NNE--  
0  --NNE--  
0  --NNE--  
0  --NNE--
```

```python
row = df_dev_pred.iloc[1]
' '.join(row.sentence.split()[row.beg:row.end])
```

Output:
```
'West Indian'
```

## Problem 4: Entity linking

Now that we have a method for predicting mention spans, we turn to the task of **entity linking**, which amounts to predicting the knowledge base entity that is referenced by a given mention. In our case, for each span we want to predict the Wikipedia page that this mention references.

### Task 4.1

Start by extending the generator function that you implemented in Task&nbsp;2.1 to labelled spans.

<!-- grade_id: cell-501e0ff440c81adf, points: 0, grade: false, solution: true -->
```python
def gold_mentions(df):
    """Yield the gold-standard mentions in a data frame.

    Args:
        df: A data frame.

    Yields:
        The gold-standard mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and entity label of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        yield((row[1],row[3],row[4],row[5]))
```

#### 🤞 Test your code

To test your code, you can run the following cell, which counts the spans yielded by your function when called on the development data (there should still be 5,917 unique tuples, just as in Task 2.1), and checks if some expected tuples are included in the results.

<!-- grade_id: cell-2cd2cac3e3b80e0b, points: 1, grade: true, solution: false -->
```python
mentions_dev_gold = set(gold_mentions(df_dev))
assert len(mentions_dev_gold) == 5917, "The number of unique returned quadruples should be the same as before."
assert ('0966-159', 1, 3, '--NME--') in mentions_dev_gold, "An expected tuple is not included in the results."
assert ('1094-020', 0, 1, 'Seattle_Mariners') in mentions_dev_gold, "An expected tuple is not included in the results."
success()
```

Output:
```
<IPython.core.display.HTML object>
```

```python
'''
i = 1
for element in mentions_dev_gold:
    print(element)
    i+=1
    if i >10:
        break
        '''
```

Output:
```
'\ni = 1\nfor element in mentions_dev_gold:\n    print(element)\n    i+=1\n    if i >10:\n        break\n        '
```

### Task 4.2

A naive baseline for entity linking on our data set is to link each mention span to the Wikipedia page name that we get when we join the tokens in the span by underscores, as is standard in Wikipedia page names. Suppose, for example, that a span contains the two tokens

    Jimi Hendrix

The baseline Wikipedia page name for this span would be

    Jimi_Hendrix

Implement this naive baseline and evaluate its performance.

**Here and in the remainder of this lab, you should base your entity predictions on the predicted spans that you computed in Problem&nbsp;3.**

<!-- grade_id: cell-f24799696158afee, points: 0, grade: false, solution: true -->
```python
'''
import requests
def baseline(df):
    """A naive baseline for entity linking that "predicts" Wikipedia
       page names from the tokens in the mention span.

    Arguments:
        df: A data frame.

    Yields:
        The predicted mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and the predicted entity label of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            if ent.label_!= 'ORDINAL'and ent.label_!= 'DATE' and ent.label_ != 'CARDINAL':
                
                # the following part is used to get "name_wiki".
                kb_id = ent.kb_id_
                wikidata_api_url = f'https://www.wikidata.org/w/api.php?action=wbgetentities&format=json&ids={kb_id}'
                response = requests.get(wikidata_api_url)
                data = response.json()
                name_wiki = data['entities'][kb_id]['labels']['en']['value'].replace(" ","_")

                yield((row[1],ent.start, ent.end, name_wiki))
                '''
```

Output:
```
'\nimport requests\ndef baseline(df):\n    """A naive baseline for entity linking that "predicts" Wikipedia\n       page names from the tokens in the mention span.\n\n    Arguments:\n        df: A data frame.\n\n    Yields:\n        The predicted mention spans in the specified data frame as\n        quadruples consisting of the sentence id, start position, end\n        position and the predicted entity label of each span.\n    """\n    # YOUR CODE HERE\n    for row in df.itertuples():\n        doc = nlp(row[2])\n        for ent in doc.ents:\n            if ent.label_!= \'ORDINAL\'and ent.label_!= \'DATE\' and ent.label_ != \'CARDINAL\':\n                \n                # the following part is used to get "name_wiki".\n                kb_id = ent.kb_id_\n                wikidata_api_url = f\'https://www.wikidata.org/w/api.php?action=wbgetentities&format=json&ids={kb_id}\'\n                response = requests.get(wikidata_api_url)\n                data = response.json()\n              
```

```python
def baseline(df):
    """A naive baseline for entity linking that "predicts" Wikipedia
       page names from the tokens in the mention span.

    Arguments:
        df: A data frame.

    Yields:
        The predicted mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and the predicted entity label of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        doc = nlp(row[2])
        for ent in doc.ents:
            if ent.label_!= 'ORDINAL'and ent.label_!= 'DATE' and ent.label_ != 'CARDINAL' and ent.label_ != 'MONEY' and ent.label_ != 'PERCENT' and ent.label_ != 'QUANTITY' and ent.label_ != 'TIME':
                yield((row[1],ent.start, ent.end, '_'.join(row[2].split(' ')[ent.start:ent.end]))) 
```

```python
# test=set(baseline(df_dev_pred))
```

```python
# df_dev.iloc[4]
```

```python
'''
i = 1
for element in test:
    print(element)
    i+=1
    if i >500:
        break 
        '''
```

Output:
```
'\ni = 1\nfor element in test:\n    print(element)\n    i+=1\n    if i >500:\n        break \n        '
```

#### 🤞 Test your code

Again, we can turn to the evaluation measures that we implemented in Problem&nbsp;1.  The expected precision should be around 32%, with an F1-score around 29%.

```python
'''
def error_report_2(df, spans_gold, spans_pred):
    false_pos = defaultdict(list)
    for s, b, e, c in spans_pred - spans_gold:
        false_pos[s].append((b, e, c))
    false_neg = defaultdict(list)
    for s, b, e, c in spans_gold - spans_pred:
        false_neg[s].append((b, e, c))
    for row in df.drop_duplicates('sentence_id').itertuples():
        if row.sentence_id in false_pos or row.sentence_id in false_neg:
            print('Sentence:', row.sentence)
            for b, e, c in false_pos[row.sentence_id]:
                print('  FP:', ' '.join(row.sentence.split()[b:e]),c)
            for b, e, c in false_neg[row.sentence_id]:
                print('  FN:', ' '.join(row.sentence.split()[b:e]),c)
                '''
```

Output:
```
"\ndef error_report_2(df, spans_gold, spans_pred):\n    false_pos = defaultdict(list)\n    for s, b, e, c in spans_pred - spans_gold:\n        false_pos[s].append((b, e, c))\n    false_neg = defaultdict(list)\n    for s, b, e, c in spans_gold - spans_pred:\n        false_neg[s].append((b, e, c))\n    for row in df.drop_duplicates('sentence_id').itertuples():\n        if row.sentence_id in false_pos or row.sentence_id in false_neg:\n            print('Sentence:', row.sentence)\n            for b, e, c in false_pos[row.sentence_id]:\n                print('  FP:', ' '.join(row.sentence.split()[b:e]),c)\n            for b, e, c in false_neg[row.sentence_id]:\n                print('  FN:', ' '.join(row.sentence.split()[b:e]),c)\n                "
```

```python
# error_report_2(df_dev.iloc[:100,:],mentions_dev_gold,set(baseline(df_dev_pred)))
```

<!-- grade_id: cell-e91c4a73987c75ce, points: 1, grade: true, solution: false -->
```python
scores = evaluation_scores(mentions_dev_gold, set(baseline(df_dev_pred)))
print_evaluation_scores(scores)
assert .31 < scores[0] < .32, "Precision should be between 31% and 32%."
assert .28 < scores[-1] < .29, "F1-score should be between 28% and 29%."
success()
```

Output:
```
Precision: 0.315, Recall: 0.263, F1: 0.287
<IPython.core.display.HTML object>
```

## Problem 5: Extending the training data using the knowledge base

State-of-the-art approaches to entity linking exploit information in knowledge bases. In our case, where Wikipedia is the knowledge base, one particularly useful type of information are links to other Wikipedia pages. In particular, we can interpret the anchor texts (the highlighted texts that you click on) as mentions of the entities (pages) that they link to. This allows us to harvest long lists of mention–entity pairings.

The following cell loads a data frame summarizing anchor texts and page references harvested from the first paragraphs of the English Wikipedia. The data frame also contains all entity mentions in the training data (but not the development or the test data).

```python
with bz2.open('kb.tsv.bz2', 'rt') as source:
    df_kb = pd.read_csv(source, sep='\t', quoting=csv.QUOTE_NONE)
```

To understand what information is available in this data, the following cell shows the entry for the anchor text `Sweden`.

```python
df_kb.loc[df_kb.mention == 'Sweden']
```

Output:
```
mention                                 entity      prob
17436  Sweden                                 Sweden  0.985768
17437  Sweden          Sweden_national_football_team  0.014173
17438  Sweden  Sweden_men's_national_ice_hockey_team  0.000059
```

As you can see, each row of the data frame contains a pair $(m, e)$ of a mention $m$ and an entity $e$, as well as the conditional probability $P(e|m)$ for mention $m$ referring to entity $e$. These probabilities were estimated based on the frequencies of mention–entity pairs in the knowledge base. The example shows that the anchor text &lsquo;Sweden&rsquo; is most often used to refer to the entity [Sweden](http://en.wikipedia.org/wiki/Sweden), but in a few cases also to refer to Sweden&rsquo;s national football and ice hockey teams. Note that references are sorted in decreasing order of probability, so that the most probable pairing come first.

Implement an entity linking method that resolves each mention to the most probable entity in the data frame. If the mention is not included in the data frame, you can predict the generic label `--NME--`.

```python
df_kb.loc[df_kb.mention=="1972 Munich Games"]
```

Output:
```
mention                entity  prob
47  1972 Munich Games  1972_Summer_Olympics   1.0
```

```python
def most_probable_method(df, df_kb):
    """An entity linker that resolves each mention to the most probably entity in a knowledge base.

    Arguments:
        df: A data frame containing the mention spans.
        df_kb: A data frame containing the knowledge base.

    Yields:
        The predicted mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and the predicted entity label of each span.
    """
    # YOUR CODE HERE
    for row in df.itertuples():
        mention = ' '.join(row.sentence.split()[row.beg:row.end])
        if len(df_kb.loc[df_kb.mention == mention]) == 0:
            label = '--NME--'
        else:
           label = df_kb.loc[df_kb.mention == mention].entity.iloc[0]
        yield(row.sentence_id,row.beg,row.end,label)
                
```

### 🤞 Test your code

We run the same evaluation as before. The expected precision should now be above 65%, with an F1-score just around 60%.

<!-- grade_id: cell-53f6994f788bb86e, points: 1, grade: true, solution: false -->
```python
scores = evaluation_scores(mentions_dev_gold, set(most_probable_method(df_dev_pred, df_kb)))
print_evaluation_scores(scores)
assert scores[0] > .65, "Precision should be above 65%."
assert .59 < scores[-1] < .61, "F1-score should be around 60%."
success()
```

Output:
```
Precision: 0.659, Recall: 0.549, F1: 0.599
<IPython.core.display.HTML object>
```

## Problem 6: Context-sensitive disambiguation

Consider the entity mention &lsquo;Lincoln&rsquo;. The most probable entity for this mention turns out to be [Lincoln, Nebraska](http://en.wikipedia.org/Lincoln,_Nebraska); but in pages about American history, we would be better off to predict [Abraham Lincoln](http://en.wikipedia.org/Abraham_Lincoln). This suggests that we should try to disambiguate between different entity references based on the textual context on the page from which the mention was taken. Your task in this last problem is to implement this idea.

Set up a dictionary that contains, for each mention $m$ that can refer to more than one entity $e$, a separate Naive Bayes classifier that is trained to predict the correct entity $e$, given the textual context of the mention. As the prior probabilities of the classifier, choose the probabilities $P(e|m)$ that you used in Problem&nbsp;5. To let you estimate the context-specific probabilities, we have compiled a data set with mention contexts:

```python
with bz2.open('contexts.tsv.bz2') as source:
    df_contexts = pd.read_csv(source, sep='\t', quoting=csv.QUOTE_NONE)
```

This data frame contains, for each ambiguous mention $m$ and each knowledge base entity $e$ to which this mention can refer, up to 100 randomly selected contexts in which $m$ is used to refer to $e$. For this data, a **context** is defined as the 5 tokens to the left and the 5 tokens to the right of the mention. Here are a few examples:

```python
df_contexts.head()
```

Output:
```
mention                 entity   
0            1970  UEFA_Champions_League  \
1            1970         FIFA_World_Cup   
2  1990 World Cup    1990_FIFA_World_Cup   
3  1990 World Cup    1990_FIFA_World_Cup   
4  1990 World Cup    1990_FIFA_World_Cup   

                                             context  
0    Cup twice the first in @ and the second in 1983  
1  America 1975 and during the @ and 1978 World C...  
2                 Manolo represented Spain at the @   
3  Hašek represented Czechoslovakia at the @ and ...  
4  renovations in 1989 for the @ The present capa...
```

Note that, in each context, the position of the mention is indicated by the `@` symbol.

From this data frame, it is easy to select the data that you need to train the classifiers – the contexts and corresponding entities for all mentions. To illustrate this, the following cell shows how to select all contexts that belong to the mention &lsquo;Lincoln&rsquo;:

```python
df_contexts.context[df_contexts.mention == 'Lincoln']
```

Output:
```
41465    Nebraska Concealed Handgun Permit In @ municip...
41466    Lazlo restaurants are located in @ and Omaha C...
41467    California Washington Overland Park Kansas @ N...
41468    City Missouri Omaha Nebraska and @ Nebraska It...
41469    by Sandhills Publishing Company in @ Nebraska USA
                               ...                        
41609                                      @ Leyton Orient
41610                    English division three Swansea @ 
41611    league membership narrowly edging out @ on goa...
41612                                          @ Cambridge
41613                                                   @ 
Name: context, Length: 149, dtype: object
```

```python
df_contexts[df_contexts.mention == 'Sweden']
```

Output:
```
mention                                 entity   
70739  Sweden                                 Sweden  \
70740  Sweden                                 Sweden   
70741  Sweden                                 Sweden   
70742  Sweden                                 Sweden   
70743  Sweden                                 Sweden   
...       ...                                    ...   
70935  Sweden          Sweden_national_football_team   
70936  Sweden          Sweden_national_football_team   
70937  Sweden          Sweden_national_football_team   
70938  Sweden          Sweden_national_football_team   
70939  Sweden  Sweden_men's_national_ice_hockey_team   

                                                 context  
70739  stayed with Nash at Oerkelljunga @ with Prem Z...  
70740  Ekerö Municipality Stockholm County southeaste...  
70741  foreign relations between Argentina and @ Both...  
70742  Pettersson July 1903 in Älvsbyn @ 18 November ...  
70743  structure located in Uddevalla 
```

Implement the context-sensitive disambiguation method and evaluate its performance.  Do this in two parts, first implementing a function that builds the classifiers _(refer to the text above for a detailed description)_, then implementing a prediction function that uses these classifiers to perform the entity prediction.

Here are some more **hints** that may help you along the way:

1. The prior probabilities for a Naive Bayes classifier can be specified using the `class_prior` option. You will have to provide the probabilities in the same order as the alphabetically sorted class (entity) names.

2. Not all mentions in the knowledge base are ambiguous, and therefore not all mentions have context data. If a mention has only one possible entity, pick that one. If a mention has no entity at all, predict the `--NME--` label.

```python
'''
my_dict = {'key1': 'value1', 'key2': 'value2', 'key3': 'value3'}
additional_dict = {'key4': 'value4', 'key5': 'value5'}

my_dict.update(additional_dict)

print(my_dict)

'''
```

Output:
```
"\nmy_dict = {'key1': 'value1', 'key2': 'value2', 'key3': 'value3'}\nadditional_dict = {'key4': 'value4', 'key5': 'value5'}\n\nmy_dict.update(additional_dict)\n\nprint(my_dict)\n\n"
```

<!-- grade_id: cell-22e330f1b6937c8d, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import make_pipeline
from sklearn.feature_extraction.text import CountVectorizer
from sklearn.naive_bayes import MultinomialNB

def build_entity_classifiers(df_kb, df_contexts):
    """Build Naive Bayes classifiers for entity prediction.

    Arguments:
        df_kb: A data frame with the knowledge base.
        df_contexts: A data frame with contexts for each mention.

    Returns:
        A dictionary where the keys are mentions and the values are Naive Bayes
        classifiers trained to predict the correct entity, given the textual
        context of the mention (as described in detail above).
    """
    # YOUR CODE HERE
    classifier = {}
    
    
    mentions_list = df_contexts.mention.unique()
    
    for mention in mentions_list:
        data_train = df_contexts[df_contexts['mention'] == mention][['entity','context']]
        pr_prior = df_kb.loc[df_kb.mention == mention][['entity','prob']].sort_values(by='entity').prob

        clf = make_pipeline(CountVectorizer(),
                            MultinomialNB(class_prior=pr_prior))
        clf.fit(data_train.context,data_train.entity)
        classifier.update({mention:clf})
    
    return(classifier) 
           
```

```python
df_contexts[df_contexts['mention'] == 'Sweden'][['entity','context']]
```

Output:
```
entity   
70739                                 Sweden  \
70740                                 Sweden   
70741                                 Sweden   
70742                                 Sweden   
70743                                 Sweden   
...                                      ...   
70935          Sweden_national_football_team   
70936          Sweden_national_football_team   
70937          Sweden_national_football_team   
70938          Sweden_national_football_team   
70939  Sweden_men's_national_ice_hockey_team   

                                                 context  
70739  stayed with Nash at Oerkelljunga @ with Prem Z...  
70740  Ekerö Municipality Stockholm County southeaste...  
70741  foreign relations between Argentina and @ Both...  
70742  Pettersson July 1903 in Älvsbyn @ 18 November ...  
70743  structure located in Uddevalla Bohuslän @ It i...  
...                                                  ...  
70935  appearance on November 13 against @ in ho
```

<!-- grade_id: cell-045c73fff60f8fbf, points: 0, grade: false, solution: true -->
```python
def extended_dictionary_method(df, classifiers, df_kb):
    """An entity linker that resolves each mention to the most probably entity in a knowledge base.

    Arguments:
        df: A data frame containing the mention spans.
        classifiers: A dictionary of classifiers as produced by the
            `build_entity_classifiers` function.
        df_kb: A data frame with the knowledge base. (Should be used
            to look up a mention if it doesn't have a classifier.)

    Yields:
        The predicted mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and the predicted entity label of each span.
    """
    # YOUR CODE HERE
    set_intersection = set(df_kb.mention).intersection(set(df_contexts.mention))

    for row in df.itertuples():
        idx, id, sentence, beg, end, label = row
        mention = ' '.join(sentence.split()[beg:end])
        if mention in set(df_kb.mention):
            if mention in set_intersection:
                classifier = classifiers[mention]
                label = classifier.predict([' '.join(sentence.split()[beg-5:end+5])])[0]
                
                
            else:
                label = df_kb.loc[df_kb.mention==mention,'entity'].iloc[0]
                
        else:
            label ='--NME--'
        
        yield(id,beg,end,label)




```

### 🤞 Test your code

The cell below shows how your functions should all come together.

```python
def extended_dictionary_method(df, classifiers, df_kb):
    """An entity linker that resolves each mention to the most probably entity in a knowledge base.

    Arguments:
        df: A data frame containing the mention spans.
        classifiers: A dictionary of classifiers as produced by the
            `build_entity_classifiers` function.
        df_kb: A data frame with the knowledge base. (Should be used
            to look up a mention if it doesn't have a classifier.)

    Yields:
        The predicted mention spans in the specified data frame as
        quadruples consisting of the sentence id, start position, end
        position and the predicted entity label of each span.
    """
    for row in df.itertuples():
        idx,id,sentence,beg,end,label = row
        mention = ' '.join(sentence.split()[beg:end])
        try:
            classfier = classifiers[mention]
            tokens = ' '.join(sentence.split())
            label = classfier.predict(tokens[beg-5:beg]+['@']+tokens[end:end+5])[0]
        except:
            try:
                label = df_kb.loc[df_kb.mention == mention].entity.iloc[0]
            except:
                label = "--NME--"


        yield (id,beg,end,label)


    # YOUR CODE HERE
    # raise NotImplementedError()
```

<!-- grade_id: cell-dda4aa4e70cd83fb, points: 1, grade: true, solution: false -->
```python
classifiers = build_entity_classifiers(df_kb, df_contexts)
mentions_dev_pred_dictionary = set(extended_dictionary_method(df_dev_pred, classifiers, df_kb))
```

Finally, the cell below evaluates the results as before. You should expect to see a small (around 1&nbsp;unit) increase in each of precision, recall, and F1.

```python
text_test= df_contexts.context[df_contexts.mention == "Swedish"][70940]
```

```python
classifiers["Swedish"].predict([text_test])[0]
```

<!-- grade_id: cell-ddde598c1d1aff97, points: 1, grade: true, solution: false -->
```python
scores = evaluation_scores(mentions_dev_gold, mentions_dev_pred_dictionary)
print_evaluation_scores(scores)
assert scores[0] > .67, "Precision should be above 67%."
assert scores[-1] > .61, "F1-score should be above 61%."
success()
```

Output:
```
Precision: 0.659, Recall: 0.549, F1: 0.599
AssertionError                            Traceback (most recent call last)
c:\Users\yj313\Desktop\LEARNING_ MATERIAL\Text Mining\2. Lab Sessions\l4\l4\TM-Lab4.ipynb Cell 96 line 3
      <a href='vscode-notebook-cell:/c%3A/Users/yj313/Desktop/LEARNING_%20MATERIAL/Text%20Mining/2.%20Lab%20Sessions/l4/l4/TM-Lab4.ipynb#Y164sZmlsZQ%3D%3D?line=0'>1</a> scores = evaluation_scores(mentions_dev_gold, mentions_dev_pred_dictionary)
      <a href='vscode-notebook-cell:/c%3A/Users/yj313/Desktop/LEARNING_%20MATERIAL/Text%20Mining/2.%20Lab%20Sessions/l4/l4/TM-Lab4.ipynb#Y164sZmlsZQ%3D%3D?line=1'>2</a> print_evaluation_scores(scores)
----> <a href='vscode-notebook-cell:/c%3A/Users/yj313/Desktop/LEARNING_%20MATERIAL/Text%20Mining/2.%20Lab%20Sessions/l4/l4/TM-Lab4.ipynb#Y164sZmlsZQ%3D%3D?line=2'>3</a> assert scores[0] > .67, "Precision should be above 67%."
      <a href='vscode-notebook-cell:/c%3A/Users/yj313/Desktop/LEARNING_%20MATERIAL/Text%20Mining/2.%20La
```

**Congratulations on finishing this lab! 👍**

<div class="alert alert-info">
    
➡️ Don't forget to **test that everything runs as expected** before you submit!

</div>